Before you turn this lab in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

**Provide your name and any collaborators below:**

<!-- grade_id: name-collabs, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

---
# Lab 0 Intro to Scala and Chisel
>  They are not intended take a significant amount of time but rather to provide examples/practice on specific and isolated features in the language. Labs are autograded so you can get quick feedback.

### Import the necessary Chisel dependencies. 
> There will be cells like these in every lab. Make sure you run them before proceeding to bring the Chisel Library into the Jupyter Notebook scope!

```scala
import $ivy.`edu.berkeley.cs::chisel3:3.5.3`
import $ivy.`edu.berkeley.cs::chiseltest:0.5.3`
import $plugin.$ivy.`edu.berkeley.cs:chisel3-plugin_2.12.15:3.5.3`
import chisel3._
import chisel3.util._
import chiseltest._
import chiseltest.RawTester.test
import scala.language.reflectiveCalls
```

## Problem 1  - Scala conditionals
> Practice with Scala's if/else. Return `"heads"` if `flip` is `true`, else return `"tails"`

<!-- grade_id: ScalaCondPractice, points: 0, grade: false, solution: true -->
```scala
def scalaCondPractice(flip: Boolean): String = {
    // YOUR CODE HERE
    ???
}
```

<!-- grade_id: scp-test, points: 2, grade: true, solution: false -->
```scala
assert(scalaCondPractice(true) == "heads")
assert(scalaCondPractice(false) == "tails")
```

## Problem 2  - Writing a Chisel Module's IO
> Fill in the IO in the module such that it takes two 3-bit `UInt`s as input and returns a 4-bit sum as output (no truncation). 

<!-- grade_id: addTwo, points: 0, grade: false, solution: true -->
```scala
class AddTwo extends Module {
    // val io = ???
    // YOUR CODE HERE
    ???
    
    io.out := io.in1 +& io.in2
}
```

<!-- grade_id: testAddTwo, points: 3, grade: true, solution: false -->
```scala
def testAddTwo: Boolean = {
    test(new AddTwo) { c =>
        for (i <- 0 until 8) {
            for (j <- 0 until 8) {
                c.io.in1.poke(i.U)
                c.io.in2.poke(j.U)
                c.io.out.expect((i+j).U)
            }
        }
    }
    true
}

assert(testAddTwo)
```

## Problem 3  - Combinational Logic
> Assign the boolean expression: `(a OR b) AND (NOT c)` to the module's output. 

<!-- grade_id: CombLogic, points: 0, grade: false, solution: true -->
```scala
class CombLogic extends Module {
    val io = IO(new Bundle {
        val a   = Input(Bool())
        val b   = Input(Bool())
        val c   = Input(Bool())
        val out = Output(Bool())
    })
    
    // YOUR CODE HERE
    ???
    
    // We can print state like this everytime `step()` is called in our test
    printf(p"a: ${io.a}, b: ${io.b}, c: ${io.c}, out: ${io.out}\n")
}
```

## Problem 4  - Combinational Logic Test
> Write your own test that tests `CombLogic` exhaustively for all input values `a, b, and c`. The module should return `true` if and only if all calls to `dut.io.out.expect(...)` succeed.

<!-- grade_id: testCombLogic, points: 0, grade: false, solution: true -->
```scala
def testCombLogic: Boolean = {
    test(new CombLogic) { dut =>
        
        // YOUR CODE HERE
        ???
    }
    true
}
```

<!-- grade_id: testCombLogic-test, points: 4, grade: true, solution: false -->
```scala
assert(testCombLogic)

```

## Problem 5  - Scala Conditional in Chisel modules
> At hardware ellaboration time, we can use Scala conditionals to change which hardware is created within a module. Implement the module such that if the `useAnd` argument is `true`, the generated hardware produces `a && b`, and otherwise produces `a || b`. The generated hardware should contain only `AND` logic or `OR` logic, but not both.

<!-- grade_id: AndOrCompileTime, points: 0, grade: false, solution: true -->
```scala
class AndOrGenerationTime(useAnd: Boolean) extends Module {
    val io = IO(new Bundle {
        val a   = Input(Bool())
        val b   = Input(Bool())
        val out = Output(Bool())
    })
    
    // YOUR CODE HERE
    ???
}
```

<!-- grade_id: testAndOrCompileTime, points: 3, grade: true, solution: false -->
```scala
def testAndOrGenerationTime(useAnd: Boolean): Boolean = {
    test(new AndOrGenerationTime(useAnd)) { dut =>
        for (a <- Seq(true, false)) {
            for (b <- Seq(true, false)) {
                dut.io.a.poke(a.B)
                dut.io.b.poke(b.B)
                if (useAnd) dut.io.out.expect((a && b).B)
                else        dut.io.out.expect((a || b).B)
            }
        }
    }
    true
}
assert(testAndOrGenerationTime(useAnd = true))
assert(testAndOrGenerationTime(useAnd = false))
```

## Problem 6  - Chisel Conditional in Chisel modules
> Generated hardware can use conditionals (i.e `Mux` or `when/elsewhen/otherwise`) to select signals. In this exercise, `useAnd` is an `Input` to the module. If `useAnd` is `true`, then the output `out` should be `a && b`, otherwise `a || b`. In this problem, both the logic for `a && b` and `a || b` hardware should be generated. You may use either a Chisel `Mux` or the Chisel `when` statement.

<!-- grade_id: AndOrRunTime, points: 0, grade: false, solution: true -->
```scala
class AndOrRunTime extends Module {
    val io = IO(new Bundle {
        val a      = Input(Bool())
        val b      = Input(Bool())
        val useAnd = Input(Bool())
        val out    = Output(Bool())
    })
    
    // YOUR CODE HERE
    ???
}
```

## Problem 7  - Last connect semantics
> When connecting Chisel components, the last connection made is the one that "wins" (exists in the generated hardware). In the module below, the default output is `5.U` because `out` is connected to `5.U` after `4.U`. Use a `when` statement to conditionally connect `8.U` to the output when the input `update` is set high, or keep the default connection when `update` is set low.

<!-- grade_id: LastConnect, points: 0, grade: false, solution: true -->
```scala
class LastConnect extends Module {
    val io = IO(new Bundle {
        val update = Input(Bool())
        val out    = Output(UInt())
    })
    
    io.out := 4.U
    io.out := 5.U
    // YOUR CODE HERE
    ???
    
}
```

<!-- grade_id: testLastConnect, points: 2, grade: true, solution: false -->
```scala
def testLastConnect: Boolean = {
    test(new LastConnect) { dut =>
        dut.io.update.poke(true.B)
        dut.io.out.expect(8.U)
        
        dut.io.update.poke(false.B)
        dut.io.out.expect(5.U)
    }
    true
}
assert(testLastConnect)
```

```scala

```

## Problem 8  - State
> We can use registers to store state across multiple cycles. In Chisel, to build registers we use `Reg` as well as `RegNext` and `RegInit` ([more info](https://www.chisel-lang.org/chisel3/docs/explanations/sequential-circuits.html)). Fill in the `Delay2` module such that `out` signal is equal to the `in` signal delayed by 2 cycles.

<!-- grade_id: delay2, points: 0, grade: false, solution: true -->
```scala
class Delay2 extends Module {
    val io = IO(new Bundle {
        val in  = Input(UInt(5.W))
        val out = Output(UInt(5.W))
    })
    
    // YOUR CODE HERE
    ???
}
```

<!-- grade_id: test-delay2, points: 2, grade: true, solution: false -->
```scala
def testDelay2: Boolean = {
    test(new Delay2) { dut =>
        // Cycle0
        dut.io.in.poke(5.U)
        dut.io.out.expect(0.U)
        dut.clock.step(1)

        // Cycle1
        dut.io.in.poke(4.U)
        dut.io.out.expect(0.U)
        dut.clock.step(1)

        // Cycle2
        dut.io.in.poke(3.U)
        dut.io.out.expect(5.U)
        dut.clock.step(1)

        // Cycle3
        dut.io.out.expect(4.U)
        dut.clock.step(1)

        // Cycle4
        dut.io.out.expect(3.U)
        dut.clock.step(1)
    }
    true
}
assert(testDelay2)
```

## Problem 9  - Accumulator
> Let's build an _accumulator_. Each cycle `en` is high, it will add `in` to it's internal total. The internal total is visible as the output `out`. The internal total should initialize to 0 on reset.

<!-- grade_id: Accumulator, points: 0, grade: false, solution: true -->
```scala
class Accumulator(w: Int) extends Module {
    val io = IO(new Bundle {
        val in  = Input(UInt(w.W))
        val en  = Input(Bool())
        val out = Output(UInt(w.W))
    })
    
    // YOUR CODE HERE
    ???
    
}
```

<!-- grade_id: test-accumulator, points: 3, grade: true, solution: false -->
```scala
def testAccumulator: Boolean = {
    test(new Accumulator(4)) { dut =>
        // Cycle 0
        dut.io.in.poke(0.U)
        dut.io.en.poke(0.B)
        dut.clock.step(1)
        dut.io.out.expect(0.U)

        // Cycle 1
        dut.io.in.poke(0.U)
        dut.io.en.poke(0.B)
        dut.clock.step(1)
        dut.io.out.expect(0.U)

        // Cycle 2
        dut.io.in.poke(3.U)
        dut.io.en.poke(0.B)
        dut.clock.step(1)
        dut.io.out.expect(0.U)

        // Cycle 3
        dut.io.in.poke(3.U)
        dut.io.en.poke(1.B)
        dut.clock.step(1)
        dut.io.out.expect(3.U)

        // Cycle 3
        dut.io.in.poke(4.U)
        dut.io.en.poke(1.B)
        dut.clock.step(1)
        dut.io.out.expect(7.U)        
    }
    true
}
assert(testAccumulator)
```

# Problem 10  - Using a Vec to Implement a Lookup Table
> Vecs allow us to dynamically index into a Chisel collection during operation in the generated hardware.

> In the problem, we use a Vec to implement a _lookup table_, which is a read-only memory (ROM) that holds precomputed results. The `SquareTable` module we provide below produces the result of squaring its input. However, the generated hardware contains only the lookup table because the multiplication is done at generation time. As a first implementation, we hardcode the IO widths and table size.

> Fill in the tester `testSquareTable` to exhaustively test all of `SquareTable`'s inputs (hint: there are 32).

```scala
class SquareTable extends Module {
    val io = IO(new Bundle {
        val x = Input(UInt(5.W))
        val xSquared = Output(UInt(10.W))
    })
    // 0.U, 1.U, ..., 31.U
    val romData = Seq.tabulate(32)(i => (i*i).U(10.W))
    show(s"romData: $romData")

    val ROM: Vec[UInt] = VecInit(romData)
    io.xSquared := ROM(io.x)
}
```

<!-- grade_id: testROM, points: 0, grade: false, solution: true -->
```scala
def testSquareTable: Boolean = {
    test(new SquareTable) { dut =>
        
        // YOUR CODE HERE
        ???
    }
    true
}
```

<!-- grade_id: test-testROM, points: 3, grade: true, solution: false -->
```scala
assert(testSquareTable)
```

## Problem 11  - Histogram
Let's put Regs and Vecs together to build a `Histogram` generator. The generated hardware will count how often it sees each input. It will take a parameter `n`, and internally it will use `n` registers to count how often it has seen inputs (from 0 to n-1). Each cycle, the `x` input chooses which internal register to increment. The `out` output should provide the new total value that will be written into the register the next cycle. In other words, `out` immediately shows new value of the count associated with `x`.

<!-- grade_id: Histogram, points: 0, grade: false, solution: true -->
```scala
class Histogram(n: Int) extends Module {
    val io = IO(new Bundle {
        val x = Input(UInt(log2Ceil(n).W))
        val out = Output(UInt(5.W))
    })
    // YOUR CODE HERE
    ???
}
```

<!-- grade_id: test-histogram, points: 6, grade: true, solution: false -->
```scala
def testHistogram: Boolean = {
    test(new Histogram(3)) { dut =>
        // Cycle0
        dut.io.x.poke(1.U)
        dut.io.out.expect(1.U)
        dut.clock.step(1)

        // Cycle1
        dut.io.x.poke(1.U)
        dut.io.out.expect(2.U)
        dut.clock.step(1)

        dut.io.x.poke(0.U)
        dut.io.out.expect(1.U)
        dut.clock.step(1)

        // Cycle 3
        dut.io.x.poke(2.U)
        dut.io.out.expect(1.U)
    }
    true
}
assert(testHistogram)
```